<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

### Import Stuff

Notice that we do not import *numpy* or *scipy* neither of these packages are need for this homework. For our solutions, the only command we needed to import was `itertools.product()`

```python
import itertools
from itertools import product
```

# Sets
Read the notebook on sets before attempting these exercises

# Problem 1

De Morgan's first law states the following for any two sets $A$ and $B$
$$(A\cup B)^c = A^c\cap B^c$$

In the following two exercises we calculate $(A\cup B)^c$ in two different ways. Both functions must take $A$, $B$ and the universal set $U$ as their inputs.

## Exercise 1.1


Write the function **complement_of_union** that first determines $A\cup B$ and then evaluates the complement of this set. Output the tuple: $\begin{pmatrix}A\cup B,\, (A\cup B)^c\end{pmatrix}$.



<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
complement_of_union(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-6, 0, 1, 2, 3}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def complement_of_union(A, B, U):
    AunB = A.union(B)
    compAunB = AunB^U
    
    return(AunB,compAunB)

```

```python
A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}

complement_of_union(A,B,U)
```

Output:
```
({-6, 0, 1, 2, 3, 4, 5, 8, 9}, {-4, -3, 7, 10})
```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}
assert( complement_of_union(A, B, U) == ({-6, 0, 1, 2, 3, 4, 5, 8, 9}, {-4, -3, 7, 10})  )

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

## Exercsise 1.2

Write the function **intersection_of_complements** that first determines $A^c$ and $B^c$ and then evaluates the intersection of their complements. Output the tuple: $\begin{pmatrix}A^c, \,  A^c\cap B^c\end{pmatrix}$

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2, 3}
B = {3, -6, 2, 0}
U = {-10, -9, -8, -7, -6, 0, 1, 2, 3, 4}
intersection_of_complements(A, B, U)
```

<font  style="color:magenta"> **Output**</font>
```
({-10, -9, -8, -7, -6, 0, 4}, {-10, -9, -8, -7, 4})
```


```python
# modify this cell

def intersection_of_complements(A, B, U):
    A_comp = U.difference(A)
    B_comp = U - B
    AcIntBc = A_comp&B_comp
    
    return(A_comp,AcIntBc)
```

```python
A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}

intersection_of_complements(A, B, U)
```

Output:
```
({-6, -4, -3, 0, 7, 8, 9, 10}, {-4, -3, 7, 10})
```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {1, 2, 3, 4, 5}
B = {0, 2, -6, 5, 8, 9}
U = A|B|{-3, 7, 10, -4}
assert(  intersection_of_complements(A, B, U) == ({-6, -4, -3, 0, 7, 8, 9, 10}, {-4, -3, 7, 10})  )

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

# Problem 2

This problem illustrates a property of cartesian products of unions of two or more sets. For four sets $A$, $B$, $S$ and $T$, the following holds:

$$(A\cup B)\times(S\cup T) = (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)$$

Write the following functions to determine $(A\cup B)\times(S\cup T)$ in two different ways.


## Exercies 2.1

Write function **product_of_unions** that first determines $(A\cup B)$ and $(S\cup T)$ and then evaluates the cartesian products of these unions. Output the tuple $\begin{pmatrix}(A\cup B),\,  (A\cup B)\times(S\cup T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
product_of_unions(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({1, 2, 3},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```


```python
# modify this cell

def product_of_unions(A, B, S, T):
    AuB = A|B
    SuT = S|T
    
    pr = set({})
    for pair in product(AuB,SuT):
        pr.add((pair))
    
    return(AuB,pr)
```

```python
A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}

product_of_unions(A,B,S,T)
```

Output:
```
({5, 6}, {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})
```

<!-- grade_id: ex3, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}
assert( product_of_unions(A, B, S, T) == \
       ({5, 6}, {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})   )

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

## Exercise 2.2

Write a function **union_of_products** that first determines $(A\times S)$ and the other three cartesian products that appear on the right hand side of the identity above, then evaluates the union of these cartesian products. Output the tuple $\begin{pmatrix}(A\times S),\,  (A\times S)\cup(A\times T)\cup(B\times S)\cup(B\times T)\end{pmatrix}$.

<font  style="color:blue"> **Code**</font>
```python
A = {1, 2}
B = {1, 3}
S = {-1, 0}
T = {0, 10}
union_of_products(A, B, S, T)
```


<font  style="color:magenta"> **Output**</font>
```
({(1, -1), (1, 0), (2, -1), (2, 0)},
 {(1, -1),
  (1, 0),
  (1, 10),
  (2, -1),
  (2, 0),
  (2, 10),
  (3, -1),
  (3, 0),
  (3, 10)})
```

```python
# modify this cell
import itertools  

def union_of_products(A, B, S, T):
    AxS = set({})
    AxT = set({})
    BxS = set({})
    BxT = set({})
    
    for pair1 in product(A,S):
        AxS.add((pair1))
    for pair2 in product(A,T):
        AxT.add((pair2))
    for pair3 in product(B,S):
        BxS.add((pair3))
    for pair4 in product(B,T):
        BxT.add((pair4))
       
    unAll = AxS|AxT|BxS|BxT
    
    return (AxS,unAll)
```

```python
A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}

union_of_products(A, B, S, T)
```

Output:
```
({(5, -1), (5, 0), (5, 1)},
 {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})
```

<!-- grade_id: ex4, points: 5, grade: true, solution: false -->
```python
# Check Function

A = {5}
B = {5, 6}
S = {-1, 0, 1}
T = {1, 2}
assert( union_of_products(A, B, S, T) == \
        ({(5, -1), (5, 0), (5, 1)}, \
         {(5, -1), (5, 0), (5, 1), (5, 2), (6, -1), (6, 0), (6, 1), (6, 2)})  \
      )

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

```python







```